## Loops

Write a loop which prints all the numbers from 0 to 12, each on their own line.

<!-- grade_id: cell-342ede9f7e8e5c2e, points: 0, grade: false, solution: true -->
```python

```

Write a function `evens` which takes a parameter `x` and prints all the even numbers between 2 and `x`, excluding `x`.

<!-- grade_id: cell-b84531928ea91a34, points: 0, grade: false, solution: true -->
```python

```

Write a function `two_powers` which prints all the powers of two with exponents 0 through 9 (ie. 1 through 512.)

<!-- grade_id: cell-abe2f23d6621f4fd, points: 0, grade: false, solution: true -->
```python

```

Write a function `total` which takes two parameters `a` and `b`, and returns the sum of all numbers from `a` to `b` (including `a` and excluding `b`.)

<!-- grade_id: cell-f78dc8c085164e7d, points: 0, grade: false, solution: true -->
```python

```

Write a function `total_evens` which takes two parameters `a` and `b`, and returns the sum of all the even numbers from `a` to `b` (including `a` and excluding `b`.)

<!-- grade_id: cell-4ea03dd27d061d72, points: 0, grade: false, solution: true -->
```python

```

## Validation

<!-- grade_id: cell-11a377cf18dbad8f, points: 0, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST
from io import StringIO
import sys
from nose.tools import assert_equal
```

<!-- grade_id: cell-740643c29763f3e5, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST

temp_out = StringIO()
sys.stdout = temp_out
evens(10)
temp_out.flush()
assert_equal(temp_out.getvalue(), "2\n4\n6\n8\n")
temp_out.close()


temp_out = StringIO()
sys.stdout = temp_out
evens(4)
temp_out.flush()
assert_equal(temp_out.getvalue(), "2\n")
temp_out.close()


temp_out = StringIO()
sys.stdout = temp_out
evens(6)
temp_out.flush()
assert_equal(temp_out.getvalue(), "2\n4\n")
temp_out.close()


temp_out = StringIO()
sys.stdout = temp_out
evens(2)
temp_out.flush()
assert_equal(temp_out.getvalue(), "")
temp_out.close()
```

<!-- grade_id: cell-8b3516e44517f954, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST

temp_out = StringIO()
sys.stdout = temp_out
two_powers()
temp_out.flush()
assert_equal(temp_out.getvalue(), "1\n2\n4\n8\n16\n32\n64\n128\n256\n512\n")
temp_out.close()
```

<!-- grade_id: cell-e537c171f9c35d50, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST

assert_equal(total(1,2), 3)
assert_equal(total(2,5), 14)
assert_equal(total(1,1), 1)
assert_equal(total(0,0), 0)
```

<!-- grade_id: cell-bb4928a437aa1622, points: 1, grade: true, solution: false -->
```python
# Ignore this cell
# DO NOT DELETE THIS CELL
### BEGIN TEST

assert_equal(total_evens(1,2), 3)
assert_equal(total_evens(2,5), 14)
assert_equal(total_evens(1,1), 1)
assert_equal(total_evens(0,0), 0)
```