## String Examples

Write a function called `check_contents`. It should take a single string as a parameter and print to report to the user whether it was a number, all uppercase, all lowercase, or mixed case.

<!-- grade_id: cell-aa83b3ded314ba8f, points: 0, grade: true, solution: true -->
```python
def check_contents(stringA):
    if stringA.isdigit():
        print("It is number")
    elif stringA.isupper():
        print("It is uppercase")
    elif stringA.islower():
        print("It is lowercase")
    else:
        print("Mixed")
check_contents("javier")
```

Output:
```
It is lowercase
```

Write a function `capitalize_first` that takes a string `s` as a parameter and returns a new string with the same contents as `s` except with the first character capitalized.

<!-- grade_id: cell-7e971816e9e525fc, points: 0, grade: false, solution: true -->
```python
def capitalize_first(s):
        print(s.capitalize())
        
capitalize_first("vassily")
```

Output:
```
Vassily
```

Write a function `add_and` that asks for user input in a loop and builds a string from that user input, putting `" and "` between the inputs the user gives. The function should stop looping and print out the result if the user enters `"exit"`.

<!-- grade_id: cell-2ee423fc116a62be, points: 0, grade: false, solution: true -->
```python

```

Write a function `switch_case` which takes a string as a parameter and returns a new string with all uppercase letters turned to lowercase and all lowercase characters turned to uppercase. For example, `switch_case("Sofia")` should return `"sOFIA"`.

<!-- grade_id: cell-211370e4ff819302, points: 0, grade: false, solution: true -->
```python

```

## Loop Terminology

A loop has four parts: initialization, condition, body, and step/increment.

In a for loop, some of these are hidden.

`for i in range(10):
    print(i)`
    
The initialization is to set `i` to the first value in the range (0), the condition is while our range hasn't run out, the body is the `print` statement, and the step/increment is to change `i` to the next value in the range.

In a while loop, you have to explicitly write all those parts.

`i = 0
while i < 10:
    print(i)
    i+=1`
  
Initialization: `i = 0`

Condition: `i < 10`

Body: `print(i)`

Step/increment: `i+=1`





`done = False
while not done:
    s = input("Done yet? ")
    if s == "done":
        done = True`
        
Initialization: `done = False`

Condition: `not done`

Body: `s = input("Done yet? ")`

Step/increment: `if s == "done": done = True`

## Tests

<!-- grade_id: cell-9823552d72ef507c, points: 1, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(capitalize_first("a"), "A")
assert_equal(capitalize_first("word"), "Word")
assert_equal(capitalize_first("another test"), "Another test")
assert_equal(capitalize_first(""), "")
```

<!-- grade_id: cell-1f12fe0ec160d102, points: 1, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(add_and(["a", "b", "c"]), "a and b and c")
assert_equal(add_and(["a"]), "a")
assert_equal(add_and([]), "")
```

<!-- grade_id: cell-ce8a39f7eac368d2, points: 1, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(switch_case("a"), "A")
assert_equal(switch_case("WORD"), "word")
assert_equal(switch_case("Another Test"), "aNOTHER tEST")
assert_equal(switch_case(""), "")
```