## Advanced String Examples

Write a function `first_word` that takes a string `s` as a parameter and returns only the first word from `s` with no spaces before or after.

<!-- grade_id: cell-65d8181611b34ec0, points: 0, grade: false, solution: true -->
```python
s.replace()
```

Write a function `get_letters` that takes a string `s` as a parameter and returns a new string containing only the alphabetical characters from `s`.

<!-- grade_id: cell-5573355e89edc511, points: 0, grade: false, solution: true -->
```python

```

Write a function `reformat_date` that changes a string with a date formatted as `09/10/2018` to a string with the date formatted as `10.09.2018`.

<!-- grade_id: cell-3ca508cc9160f77b, points: 0, grade: false, solution: true -->
```python

```

Write a function `alternate_caps` that takes a string `s` as a parameter and creates a new string that has the first word from `s` but all uppercase, the second word from `s` but all lowercase, the third word from `s` but all uppercase, and so on.

<!-- grade_id: cell-2263955949917697, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-1ffa8570694620c1, points: 0, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(first_word("a"), "a")
assert_equal(first_word("word one"), "word")
assert_equal(first_word("another test here"), "another")
assert_equal(first_word(""), "")
```

<!-- grade_id: cell-885106d2582288e2, points: 0, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(get_letters("a"), "a")
assert_equal(get_letters("word one 1"), "wordone")
assert_equal(get_letters("c4tf1sh!"), "ctfsh")
assert_equal(get_letters(""), "")
```

<!-- grade_id: cell-9bd766667203b93b, points: 0, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(reformat_date("09/10/2018"), "10.09.2018")
assert_equal(reformat_date("01/01/2018"), "01.01.2022")
```

<!-- grade_id: cell-ee863e0a921ef790, points: 0, grade: true, solution: false -->
```python
#tests for capitalize_first
from nose.tools import assert_equal

### BEGIN TEST
assert_equal(first_word("a"), "A")
assert_equal(first_word("word one"), "Word one")
assert_equal(first_word("another test here"), "Another test Here")
assert_equal(first_word(""), "")
```